```python
%%capture
### Run this cell to check your answers as you go along
import sys
!{sys.executable} -m pip install otter-grader
# Initialize Otter 
import otter
grader = otter.Notebook("group_activity4.ipynb")
```

# In-class coding exercise #4
Objective: In this exercise you will practice reading data from a text file using **Numpy** `genfromtext` and plotting it using best practices.

## Introduction to group coding exercises
Today you’ll work on this exercise in new groups of 3-4, submitting a single notebook file at the end of the class period. Decide amongst yourselves which member will upload the completed notebook to Gradescope this week. Make sure that everyone takes a turn being the “Uploader”. _You cannot upload the final code two weeks in a row._

Designate a different group member to be the "Reporter". The Reporter will be in charge of participating in the group discussion at the end of the class session.

### Workflow
Each question will be timed to ensure that everyone gets to work on at least a part of every question. Group activities are not graded by completeness or correctness, but by effort. We will be breaking down each question in the following order:  
1. Independent work 
2. Group work and discussion on coding question
3. Group work and discussion on reflection questions

You are welcome and encouraged to communicate with other groups and the teaching team when you feel stuck on a problem. 

As a reminder, we will be grading based best practices in coding. These include: 
1) Variables are used to store objects

2) Code is commented adequately

3) Variables are names appropriately

4) Code is efficient with minimal unnecessary lines 

5) Documenting help from outside sources, such as from other groups or online documentation. 

6) Final notebook fully runs from start to finish. A good way to check this is restarting the kernel and fully running through all the cells to check for any errors.

### Storing your answers
In the code cells where you will write your answers, there will be comments denoting:

"**# your code**"

and 

"**# answer variables**"

You may store any intermediary variables in the **your code** section. If you do not have any intermediary, you can also store your answer directly in the answer variables.

## Note here **and in the Gradescope submission** each of your group members:
1.
2. 
3.

# Question #1: Accessing data (35 minutes)

Sea level rise is a major concern for many coastal cities, such as Seattle. NOAA tracks the monthly average sea level in several monitoring stations. The seasonal, or interannual changes, can also be removed to look at the large-scale trend. Today, we will download and plot the monthly mean sea levels in a Seattle monitoring station using NOAA data. 

![Trends from 1899 to 2022](https://tidesandcurrents.noaa.gov/sltrends/plots/9447130_meantrend.png)

## Instructions
1. Download data from the [NOAA tides and currents website](https://tidesandcurrents.noaa.gov/sltrends/sltrends_station.shtml?id=9447130) by clickin the blue **Export to TXT** button below the figure. Download the data to your local computer by typing `ctrl + s / command + s` depending on your machine. 


2. Upload the .txt file from your local computer to the same directory as this notebook in your JupyterHub. Rename it to "noaa_seattle_sea_levels.txt". Make sure you are not duplicating the ".txt" extension at the end!


3. In the cell below, store your file name as a string in the `filepath` variable, and use it with the `open()` function to create a file object in the `file_obj` variable.


4. Find how many **header lines** are in your file using `readline( )` and a for loop. Store the number of lines you found in the `header_lines` variable. You may use either a for loop or a while loop. You will need to count how many lines are in the header from the print output, and you may "hard code" your answer in the `header_lines` variable for this part.

_HINT: You can loop as many or as little times as you want, as long as you can find the correct number of header lines._ 



_Store your solutions in the cell below_

```python
# your code

# answer variables
filepath = ""
file_obj = ""
header_lines= ""
```

```python
grader.check("Question 1")
```

# Question #2: loading and plotting data (40 minutes)

## Instructions
1. Using `np.genfromtxt`, extract the monthly mean sea level column ("monthly_msl") from the file you just downloaded (column 2). Store the resulting array in the `'msl` variable and print it.
2. Using the numpy `arange()` function, create an array with the _same length_ as the sea level data. Store the array in the `months` variable. This array will represent the number of months of data collection (starting from 0).
3. Use matplotlib `subplots()` function to make a plot of the data. Customize the following such that they are different from the default settings (be creative!):
> 1. Line color
> 2. Line width
> 3. Line style
> 3. Marker shape
> 4. Marker color

4. **Make sure your plot has all of the required plot elements, including:**
> 1. x and y axis labels with units
> 2. a meaningful, descriptive title
> 3. font size above 12

5. Answer the following questions in a markdown cell below:
> 1. What does this plot tell you about the sea level in Seattle? For example, does it look like it's increasing? Does this seem a lot to you?
> 2. Why do you think some months have very high or very low sea levels compared to other months? Do you think this is normal?
> 3. If you were to present this figure to a group of policymakers in your city, what are some key aspects of you would focus on? Why?

If you are unsure of how a function works, try looking for it in the documentation first:
- https://numpy.org/doc/stable/reference/index.html
- https://matplotlib.org/3.3.2/index.html


#### Type your answer in the cell below:

```python
# your code

# answer variables
msl = ""
months = ""

# print your answers below 

# create your plot 

```

```python
grader.check("Question 2")
```

<!-- BEGIN QUESTION -->

# Reflection questions (5 minutes)
The purpose of the reflection is to inform us as instructors about students comfort level with course content. We use these answers to inform how we spend class time and design coursework in subsequent weeks. This question is graded for completeness, so please answer each question in the text box below. Be concise in your answers (max. 2 sentences). 

1) What do you feel you excelled at in this exercise? Why?

2) What did you struggle with most in the exercise? Why?

3) Is there any section of the question that you did not complete? If so, briefly describe why and the section you spent the most time on. 

4) Is there any topic you feel we need to revisit or review in class? Why?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

# Optional
Only work on these questions if you have enough time between finishing questions 1-2 and submitting the assignment! Give your team at least 10 minutes to push to GitHub and submit to Gradescope. 

## Instructions
1. numpy `arange()` creates a nice array to represent a proxy for time. But, it is still hard to see the differences between months across years. Use numpy `genfromtxt()` to read in columns (0,1) as dtype string to get the year and month from `filepath`.

2. Use a loop to convert each [year, month] array into a datetime object using strptime(). Don't forget to import datetime. Think about how you can convert your [year, month] arrays into 1 string to pass into strptime(). _HINT: Use the_ `join()` _function._ Store your converted datetimes in a new list or numpy array.

3. Plot your new datetime array on your figure. Notice how matplotlib only plots every 20 years.
> 1. Use this [StackOverflow](https://stackoverflow.com/questions/25538520/change-tick-frequency-on-x-time-not-number-frequency) page and check the first answer to help you import the matplotlib.dates class correctly. 
> 2. Change the x-tick frequency to every 5 years. Take note of the last 2 lines in the answer and how they used `ax.xaxis.set_major_locator()` and `ax.xaxis.set_major_formatter()`. Use `YearLocator()` instead of `MinuteLocator()` to update your frequency. The [documentation](https://matplotlib.org/stable/api/dates_api.html#matplotlib.dates.YearLocator) will help.
> 3. Rotate your x-axis so the labels don't run over each other (look at example 3 on the [documentation](https://matplotlib.org/stable/gallery/ticks/ticks_too_many.html) for help). Try to get your x-axis to look as close as possible to the example below. You do not have to include the trend and legend if you do not get to Part 4.

4. Read in the linear trend from `filepath` (column 3) and plot it on top of the mean monthly data in a figure. Try to get your plot to look as close to the example below, and don't forget to include a legend! What does the trend tell us about our data? Answer the question in a comment below your figure.

![Final Plot](https://i.ibb.co/CsMpmSw/Screen-Shot-2024-10-23-at-3-57-39-PM.png)

```python
# your code

# your plot

```

<!-- END QUESTION -->



Tests for Question 1 (points: 0)
```
>>> header_lines == 6
True
```
failure_message: Check your header_lines!
```
>>> filepath=='noaa_seattle_sea_levels.txt'
True
```
failure_message: Did you correctly rename your file in filepath?


Tests for Question 2 (points: 0)
```
>>> all(msl==np.genfromtxt(filepath, skip_header=6, usecols=2, dtype='float'))
True
```
failure_message: Check your msl variable! Are you skipping the right amount of header lines?
```
>>> len(months)==len(msl)
True
```
failure_message: Check that your months variable is the same length as msl!
